# RO47019: Intelligent Control Systems Practical Assignment
* Period: 2024-2025, Q4
* Course homepage: https://brightspace.tudelft.nl/d2l/home/682445
* Instructor: Cosimo Della Santina (C.DellaSantina@tudelft.nl)
* Teaching assistant: Niels Stienen (N.L.Stienen@student.tudelft.nl)
* (c) TU Delft, 2025

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE` and remove `raise NotImplementedError()` afterwards. Moreover, if you see an empty cell, please **do not** delete it, instead run that cell as you would run all other cells. Finally, please **do not** add any extra cells to this notebook or change the existing cells unless you are explicitly asked to do so.

Please fill in your name(s) and other required details below:

```python
# Please fill in your names, student numbers, netID, and emails below.
STUDENT_1_NAME = ""
STUDENT_1_STUDENT_NUMBER = ""
STUDENT_1_NETID = ""
STUDENT_1_EMAIL = ""
```

<!-- grade_id: cell-9f148ec62e0de49c, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled
assert STUDENT_1_NAME != ""
assert STUDENT_1_STUDENT_NUMBER != ""
assert STUDENT_1_NETID != ""
assert STUDENT_1_EMAIL != ""
```

### General announcements

* Do *not* share your solutions (also after the course is finished), and do *not* copy solutions from others. By submitting your solutions, you claim that you alone are responsible for this code.

* Please post your questions regarding this assignment in the correct support forum on Brightspace, this way everybody can benefit from the response. Please note that it is **not** allowed to post any code relating to solution attempts. If you do have a particular question that you want to ask directly, please use the scheduled Q&A hours to ask the TA or if not possible otherwise, send an email to the instructor or TA.

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code! These lines are just there as a reminder for you that you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore.

* This [Jupyter notebook](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. 

* Please note that the above mentioned _read-only_ protection only works in Jupyter Notebook, and it does not work if you open this notebook in another editor (e.g., VSCode, PyCharm, etc.). Therefore, we recommend that you only use Jupyter Notebook for this course. If you use any other editor, you may accidentally delete cells, modify the tests, etc., which would cause you to lose points.

* If you edit a function that is imported in another notebook, you need to **restart the kernel** of the notebook where you are using the function. Otherwise, the changes will not be effective.

* **IMPORTANT**: Please make sure that your code executes without any errors before submitting the notebook. An easy way to ensure this is to use the validation script as described in the README.

```python

```

# Classical model-based controllers for mechanical systems
This file should solely contain the functions implementing the various feedback and feedforward controllers. After the implementing the controllers here, please run the Jupyter notebooks dedicated to each task for simulating the closed-loop system. In particular, this includes for Task 2.a:

- `task_2a-1_pd_control.ipynb`
- `task_2a-2_pd_plus_gravity_compensation_control.ipynb`
- `task_2a-3_pd_plus_feedforward_control.ipynb`
- `task_2a-4_pd_plus_control.ipynb`

```python
# import dependencies
from functools import partial
from jax import Array, jit, random
from jax import numpy as jnp
from typing import Callable
```

<!-- grade_id: A-ctrl_fb_pd, points: 0, grade: false, solution: true -->
```python
@jit
def ctrl_fb_pd(
    th: Array,
    th_d: Array,
    th_des: Array,
    th_d_des: Array,
    kp: Array = jnp.zeros((2, 2)),
    kd: Array = jnp.zeros((2, 2)),
):
    """
    PD feedback controller applied to error between target and actual absolute link angles.
    Args:
         th: link angles [rad] of shape (2, )
         th_d: link velocities [rad/s] of shape (2, )
         th_des: desired link angles [rad] of shape (2, )
         th_d_des: desired link velocities [rad/s] of shape (2, )
         kp: proportional gains of shape (2, 2)
         kd: derivative gains of shape (2, 2)
     Returns:
         tau_fb: Link torques computed by the PD controller [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_fb
```

<!-- grade_id: A-ctrl_fb_pd_rel, points: 0, grade: false, solution: true -->
```python
@jit
def ctrl_fb_pd_rel(
    th: Array,
    th_d: Array,
    th_des: Array,
    th_d_des: Array,
    kp: Array = jnp.zeros((2, 2)),
    kd: Array = jnp.zeros((2, 2)),
):
    """
    PD feedback controller applied to relative angles (i.e., joint angles)
    Args:
         th: link angles [rad] of shape (2, )
         th_d: link velocities [rad/s] of shape (2, )
         th_des: desired link angles [rad] of shape (2, )
         th_d_des: desired link velocities [rad/s] of shape (2, )
         kp: proportional gains of shape (2, 2)
         kd: derivative gains of shape (2, 2)
     Returns:
         tau_fb_rel: Link torques computed by the PD controller [Nm] of shape (2, )
    """
    # compute the relative angles and velocities
    # th_rel =
    # th_d_rel =
    # th_des_rel =
    # th_d_des_rel =
    # YOUR CODE HERE
    raise NotImplementedError()

    # compute the PD torque based on the error of the joint angles and velocities
    # tau_fb_rel =
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_fb_rel
```

<!-- grade_id: A-ctr_ff_gravity_compensation, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_ff_gravity_compensation(
    dynamical_matrices_fn: Callable,
    th: Array,
    th_d: Array,
    th_des: Array,
    th_d_des: Array,
    th_dd_des: Array,
) -> Array:
    """
    Computes the feed-forward control term for gravity compensation
    Chapter 8.5.1 / Eq. (8.50) in "Robotics: Modelling, Planning and Control" by Siciliano et al.
    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )

    Returns:
        tau_ff: feed-forward control torques [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_ff
```

<!-- grade_id: A-ctr_ff_pure_feedforward, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_ff_feedforward(
    dynamical_matrices_fn: Callable,
    th: Array,
    th_d: Array,
    th_des: Array,
    th_d_des: Array,
    th_dd_des: Array,
) -> Array:
    """
    Computes the control actions of a feedforward controller, which solely relies on the
    desired link angular accelerations, velocities and angles.

    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )

    Returns:
        tau_ff: feed-forward control torques [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_ff
```

<!-- grade_id: A-ctr_ff_pd_plus, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_ff_pd_plus(
    dynamical_matrices_fn: Callable,
    th: Array,
    th_d: Array,
    th_des: Array,
    th_d_des: Array,
    th_dd_des: Array,
) -> Array:
    """
    Computes the mixed feed-forward control term for PD+ control,
    which compensates the dynamics by evaluating the model at the current system state.
    This corresponds to the Paden and Panja Scheme, as described in Chapter 7.3.4 / E. (7.62) of
    "Brogliato, Bernard, et al. "Dissipative systems analysis and control." Theory and Applications 2 (2007): 2-5."
    for the trajectory tracking case.

    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )

    Returns:
        tau_ff: feed-forward control torques [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_ff
```

```python

```