# RO47019: Intelligent Control Systems Practical Assignment
* Period: 2022-2023, Q3
* Course homepage: https://brightspace.tudelft.nl/d2l/home/500969
* Instructor: Cosimo Della Santina (C.DellaSantina@tudelft.nl)
* Teaching assistant: Ruben Martin Rodriguez (R.MartinRodriguez@student.tudelft.nl)
* (c) TU Delft, 2023

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Remove `raise NotImplementedError()` afterwards. Moreover, if you see an empty cell, please DO NOT delete it, instead run that cell as you would run all other cells. Please fill in your name(s) and other required details below:

```python
# Please fill in your names, student numbers, netID, and emails below.
STUDENT_1_NAME = ""
STUDENT_1_STUDENT_NUMBER = ""
STUDENT_1_NETID = ""
STUDENT_1_EMAIL = ""
```

<!-- grade_id: cell-9f148ec62e0de49c, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled
assert STUDENT_1_NAME != ""
assert STUDENT_1_STUDENT_NUMBER != ""
assert STUDENT_1_NETID != ""
assert STUDENT_1_EMAIL != ""
```

### General announcements

* Do *not* share your solutions, and do *not* copy solutions from others. By submitting your solutions, you claim that you alone are responsible for this code.

* Do *not* email questions directly, since we want to provide everybody with the same information and avoid repeating the same answers. Instead, please post your questions regarding this assignment in the correct support forum on Brightspace, this way everybody can benefit from the response. If you do have a particular question that you want to ask directly, please use the scheduled Q&A hours to ask the TA.

* There is a strict deadline for each assignment. Students are responsible to ensure that they have uploaded their work in time. So, please double check that your upload succeeded to the Brightspace and avoid any late penalties.

* This [Jupyter notebook](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. You are free to add additional cells, but if you find a cell that you cannot modify or remove, please know that this is on purpose.

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code! These lines are just there as a reminder for youthat you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

```python

```

# Classical model-based controllers for mechanical systems
This file should solely contain the functions implementing the various feedback and feedforward controllers. After the implementing the controllers here, please run the Jupyter notebooks dedicated to each task for simulating the closed-loop system. In particular, this includes for Task 2.a:

- `task_2a-1_pd_control.ipynb`
- `task_2a-2_pd_plus_gravity_compensation_control.ipynb`
- `task_2a-3_pd_plus_feedforward_control.ipynb`
- `task_2a-4_pd_plus_control.ipynb`
- `task_2a-5_inverse_dynamics_control.ipynb`

```python
# import dependencies
from functools import partial
from jax.config import config as jax_config

jax_config.update("jax_platform_name", "cpu")  # set default device to 'cpu'
jax_config.update("jax_enable_x64", True)  # double precision
from jax import jit, random
from jax import numpy as jnp
from typing import Callable
```

<!-- grade_id: A-ctrl_fb_pd, points: 0, grade: false, solution: true -->
```python
@jit
def ctrl_fb_pd(
    th: jnp.ndarray,
    th_d: jnp.ndarray,
    th_des: jnp.ndarray,
    th_d_des: jnp.ndarray,
    kp: jnp.ndarray = jnp.zeros((2, 2)),
    kd: jnp.ndarray = jnp.zeros((2, 2)),
):
    """
    PD feedback controller
    Args:
         th: joint angles [rad] of shape (2, )
         th_d: joint velocities [rad/s] of shape (2, )
         th_des: desired joint angles [rad] of shape (2, )
         th_d_des: desired joint velocities [rad/s] of shape (2, )
         kp: proportional gains of shape (2, 2)
         kd: derivative gains of shape (2, 2)
     Returns:
         tau_fb: Joint torques computed by the PD controller [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_fb
```

<!-- grade_id: A-ctr_ff_gravity_compensation, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_ff_gravity_compensation(
    dynamical_matrices_fn: Callable,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
    th_des: jnp.ndarray,
    th_d_des: jnp.ndarray,
    th_dd_des: jnp.ndarray,
) -> jnp.ndarray:
    """
    Computes the feed-forward control term for gravity compensation
    Chapter 8.5.1 / Eq. (8.50) in "Robotics: Modelling, Planning and Control" by Siciliano et al.
    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )

    Returns:
        tau_ff: feed-forward control torques [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_ff
```

<!-- grade_id: A-ctr_ff_pure_feedforward, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_ff_feedforward(
    dynamical_matrices_fn: Callable,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
    th_des: jnp.ndarray,
    th_d_des: jnp.ndarray,
    th_dd_des: jnp.ndarray,
) -> jnp.ndarray:
    """
    Computes the control actions of a feedforward controller, which solely relies on the
    desired link angular accelerations, velocities and angles.

    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )

    Returns:
        tau_ff: feed-forward control torques [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_ff
```

<!-- grade_id: A-ctr_ff_pd_plus, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_ff_pd_plus(
    dynamical_matrices_fn: Callable,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
    th_des: jnp.ndarray,
    th_d_des: jnp.ndarray,
    th_dd_des: jnp.ndarray,
) -> jnp.ndarray:
    """
    Computes the mixed feed-forward control term for PD+ control,
    which compensates the dynamics by evaluating the model at the current system state.
    This corresponds to the Paden and Panja Scheme, as described in Chapter 7.3.4 / E. (7.62) of
    "Brogliato, Bernard, et al. "Dissipative systems analysis and control." Theory and Applications 2 (2007): 2-5."
    for the trajectory tracking case.

    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )

    Returns:
        tau_ff: feed-forward control torques [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau_ff
```

<!-- grade_id: A-ctr_inverse_dynamics, points: 0, grade: false, solution: true -->
```python
@partial(
    jit,
    static_argnums=(0,),
    static_argnames=("dynamical_matrices_fn",),
)
def ctrl_inverse_dynamics(
    dynamical_matrices_fn: Callable,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
    th_des: jnp.ndarray,
    th_d_des: jnp.ndarray,
    th_dd_des: jnp.ndarray,
    kp: float = jnp.zeros((2, 2)),
    kd: float = jnp.zeros((2, 2)),
) -> jnp.ndarray:
    """
    Implements an inverse dynamics / computed torque controller according to Chapter 8.5.2 of the book
    "Robotics, Modelling, Planning and Control" by Siciliano et al.

    Args:
        dynamical_matrices_fn: function that computes the dynamical matrices of the system.
            Needs to implement the following signature: M, C, G = dynamical_matrices_fn(th, th_d)
        th: current link angles [rad] of shape (2, )
        th_d: current link angular velocities [rad/s] of shape (2, )
        th_des: desired link angles [rad] of shape (2, )
        th_d_des: desired link angular velocities [rad/s] of shape (2, )
        th_dd_des: desired link angular accelerations [rad/s^2] of shape (2, )
        kp: proportional feedback gains of shape (2, 2)
        kd: derivative feedback gains of shape (2, 2)
    Returns:
        tau: control torques (feed-forward and feed-back) [Nm] of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return tau
```

```python

```