# RO47019: Intelligent Control Systems Practical Assignment
* Period: 2022-2023, Q3
* Course homepage: https://brightspace.tudelft.nl/d2l/home/500969
* Instructor: Cosimo Della Santina (C.DellaSantina@tudelft.nl)
* Teaching assistant: Ruben Martin Rodriguez (R.MartinRodriguez@student.tudelft.nl)
* (c) TU Delft, 2023

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Remove `raise NotImplementedError()` afterwards. Moreover, if you see an empty cell, please DO NOT delete it, instead run that cell as you would run all other cells. Please fill in your name(s) and other required details below:

```python
# Please fill in your names, student numbers, netID, and emails below.
STUDENT_1_NAME = ""
STUDENT_1_STUDENT_NUMBER = ""
STUDENT_1_NETID = ""
STUDENT_1_EMAIL = ""
```

<!-- grade_id: cell-9f148ec62e0de49c, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled
assert STUDENT_1_NAME != ""
assert STUDENT_1_STUDENT_NUMBER != ""
assert STUDENT_1_NETID != ""
assert STUDENT_1_EMAIL != ""
```

### General announcements

* Do *not* share your solutions, and do *not* copy solutions from others. By submitting your solutions, you claim that you alone are responsible for this code.

* Do *not* email questions directly, since we want to provide everybody with the same information and avoid repeating the same answers. Instead, please post your questions regarding this assignment in the correct support forum on Brightspace, this way everybody can benefit from the response. If you do have a particular question that you want to ask directly, please use the scheduled Q&A hours to ask the TA.

* There is a strict deadline for each assignment. Students are responsible to ensure that they have uploaded their work in time. So, please double check that your upload succeeded to the Brightspace and avoid any late penalties.

* This [Jupyter notebook](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. You are free to add additional cells, but if you find a cell that you cannot modify or remove, please know that this is on purpose.

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code! These lines are just there as a reminder for youthat you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

```python

```

# Lagrangian Neural Network (LNN)

**Author:** Maximilian Stölzle (M.W.Stolzle@tudelft.nl)

This notebook will contain functions implement a Lagrangian neural network. Please follow the notebook `task_2c-2_lnn_implementation.ipynb`, which guides you through these implementations.

```python
from flax import linen as nn
from flax.core import FrozenDict
from functools import partial
from jax.config import config as jax_config

jax_config.update("jax_platform_name", "cpu")  # set default device to 'cpu'
jax_config.update("jax_enable_x64", True)  # double precision
import jax
from jax import jit, debug
import jax.numpy as jnp
import numpy as np
from typing import Dict, Tuple

from jax_double_pendulum.integrators import rk4_step
from jax_double_pendulum.utils import normalize_link_angles
```

<!-- grade_id: cell-90677efa4e583357, points: 0, grade: false, solution: true -->
```python
class MassMatrixNN(nn.Module):
    """
    Neural network to approximate the mass matrix.
    """

    num_hidden: int = 32  # Number of hidden units per intermediate layer

    diagonal_shift = 0.001  # shifting the diagonal before activation function
    diagonal_eps = 0.002  # small value added to the resulting diagonal

    @nn.compact
    def __call__(self, th: jnp.ndarray) -> jnp.ndarray:
        """
        Evaluate the mass matrix for the current neural network parameters
        Args:
            th: link angles of shape (2, ).
                We assume that the link angles are already normalized to the interval [-pi, pi].
        Returns:
            M: mass matrix of shape (2, 2)
        """
        num_dof = th.shape[-1]  # Degrees of Freedom of the robot
        num_nn_outputs = int((num_dof**2 + num_dof) / 2)

        # implement the neural network layers
        # use nn.Dense and nn.softplus layers
        # the output of the last layer needs to be saved in the variable `m`
        # the elements of the triangular matrix are the outputs of the network
        m = jnp.zeros((num_nn_outputs,))
        # YOUR CODE HERE
        raise NotImplementedError()

        # split-off the first num_dof as the diagonal entries
        l_diagonal, l_off_diagonal = jnp.split(
            m,
            np.array(
                [
                    num_dof,
                ]
            ),
            axis=-1,
        )

        # ensure positive diagonal
        # first, add `self.diagonal_shift` to the diagonal
        # then, apply softplus
        # finally, add `self.diagonal_eps` to the diagonal
        # YOUR CODE HERE
        raise NotImplementedError()

        # Calculate the indices of the diagonal elements of L:
        indices_diag = np.arange(num_dof, dtype=int) + 1
        indices_diag = (indices_diag * (indices_diag + 1) / 2 - 1).astype(int)  # [0, 2]
        # Calculate the indices of the off-diagonal elements of L:
        indices_off_diag = np.setdiff1d(np.arange(num_nn_outputs), indices_diag)  # [1]
        # Indexing for concatenation of l_diagonal and l_off_diagonal
        indices_nn_output = np.hstack((indices_diag, indices_off_diag))  # [0, 2, 1]

        # vector of lower triangular matrix (i.e. flattened lower triangular matrix)
        vec_tril = jnp.concatenate([l_diagonal, l_off_diagonal], axis=-1)[
            ..., indices_nn_output
        ]

        # construct empty triangular matrix
        tril_mat = jnp.zeros((num_dof, num_dof))
        # (i, j) indices of lower triangular matrix
        indices_tril = np.tril_indices(num_dof)  # (array([0, 1, 1]), array([0, 0, 1]))
        # populate triangular matrix from vector
        tril_mat = tril_mat.at[indices_tril].set(vec_tril[:])

        # construct mass matrix from triangular matrix
        M = tril_mat @ tril_mat.transpose()

        return M
```

<!-- grade_id: cell-e2583bbbc2ae1715, points: 0, grade: false, solution: true -->
```python
class PotentialEnergyNN(nn.Module):
    """
    Neural network to approximate the potential energy.
    """

    num_hidden: int = 32  # Number of hidden units per intermediate layer

    @nn.compact
    def __call__(self, th: jnp.ndarray) -> jnp.ndarray:
        """
        Evaluate the potential energy for the current neural network parameters
        Args:
            th: link angles of shape (2, ).
                We assume that the link angles are already normalized to the interval [-pi, pi].
        Returns:
            U: potential energy of shape ( )
        """

        # YOUR CODE HERE
        raise NotImplementedError()

        return U
```

<!-- grade_id: cell-a63dfc33b5289c23, points: 0, grade: false, solution: true -->
```python
@jit
def kinetic_energy_fn(
    mass_matrix_nn_params: Dict, th: jnp.ndarray, th_d: jnp.ndarray
) -> jnp.ndarray:
    """
    Compute the kinetic energy of the system using a learned neural-network-based mass matrix
    Args:
        mass_matrix_nn_params: parameters of the MassMatrixNN
        th: link angles of shape (2, ).
            We assume that the link angles are already normalized to the interval [-pi, pi].
        th_d: link angular velocities of double pendulum of shape (2, )
    Returns:
        T: kinetic energy of shape ( )
    """
    # evaluate mass matrix at current system state
    M = MassMatrixNN().apply({"params": mass_matrix_nn_params}, th)

    # compute kinetic energy
    T = jnp.array(0)
    # YOUR CODE HERE
    raise NotImplementedError()

    return T
```

<!-- grade_id: cell-6a3a2cc1abe3951d, points: 0, grade: false, solution: true -->
```python
@jit
def potential_energy_fn(
    potential_energy_nn_params: FrozenDict, th: jnp.ndarray
) -> jnp.ndarray:
    """
    Compute the potential energy of the system using the `PotentialEnergyNN` neural network
    Args:
        potential_energy_nn_params: parameters of the PotentialEnergyNN
        th: link angles of shape (2, ).
            We assume that the link angles are already normalized to the interval [-pi, pi].
    Returns:
        U: potential energy of shape ( )
    """
    # evaluate the `PotentialEnergyNN` at current system state
    U = jnp.array(0)
    # YOUR CODE HERE
    raise NotImplementedError()
    return U
```

<!-- grade_id: cell-216db944065f69db, points: 0, grade: false, solution: true -->
```python
@jit
def lagrangian_fn(
    mass_matrix_nn_params: FrozenDict,
    potential_energy_nn_params: FrozenDict,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
) -> jnp.ndarray:
    """
    Compute the Lagrangian of the system using a learned neural-network-based mass matrix and potential energy
    Args:
        mass_matrix_nn_params: parameters of the MassMatrixNN
        potential_energy_nn_params: parameters of the PotentialEnergyNN
        th: link angles of shape (2, ).
            We assume that the link angles are already normalized to the interval [-pi, pi].
        th_d: link angular velocities of double pendulum of shape (2, )
    Returns:
        L: Lagrangian of shape ( )
    """
    # Compute the Lagrangian
    L = jnp.array(0)
    # YOUR CODE HERE
    raise NotImplementedError()

    return L
```

<!-- grade_id: cell-82a33f0bd7f4ddcc, points: 0, grade: false, solution: true -->
```python
@jit
def dynamical_matrices(
    mass_matrix_nn_params: FrozenDict,
    potential_energy_nn_params: FrozenDict,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
) -> Tuple[jnp.ndarray, jnp.ndarray, jnp.ndarray]:
    """
    Compute the dynamical matrices of the system using a learned neural-network-based
    mass matrix and potential energy
    The resulting Equations of Motion (EoM) are given by:
        M @ th_dd + C @ th_d + G = tau
    where tau are the external torques applied on the links.

    Args:
        mass_matrix_nn_params: parameters of the MassMatrixNN
        potential_energy_nn_params: parameters of the PotentialEnergyNN
        th: link angles of double pendulum of shape (2, )
        th_d: link angular velocities of double pendulum of shape (2, )

    Returns:
        M: mass matrix of shape (2, 2)
        C: coriolis and centrifugal matrix of shape (2, 2)
        G: gravity matrix of shape (2, )
    """
    # the neural network needs to be robust to the redundancy of th = th + n * 2*pi
    # therefore, we first project the given link angles onto the interval [-pi, pi]
    # Hint: use the `jax_double_pendulum.utils.normalize_link_angles` function.
    # Hint: use `jax.grad` or `jax.value_and_grad` for the 1st-order partial derivatives
    # Hint: use the `jax.hessian` function for computing the 2nd-order partial derivatives.
    # YOUR CODE HERE
    raise NotImplementedError()

    # Compute the dynamical matrices by taking the appropiate partial derivatives of the Lagrangian
    M, C, G = jnp.zeros((2, 2)), jnp.zeros((2, 2)), jnp.zeros((2,))
    # YOUR CODE HERE
    raise NotImplementedError()

    return M, C, G
```

<!-- grade_id: cell-2538b9eeb0fe06f9, points: 0, grade: false, solution: true -->
```python
@jit
def continuous_forward_dynamics(
    mass_matrix_nn_params: FrozenDict,
    potential_energy_nn_params: FrozenDict,
    th: jnp.ndarray,
    th_d: jnp.ndarray,
    tau: jnp.ndarray = jnp.zeros((2,)),
) -> jnp.ndarray:
    """
    Compute the continuous forward dynamics of the system using a learned neural-network-based
    mass matrix and potential energy
    Args:
        mass_matrix_nn_params: parameters of the MassMatrixNN
        potential_energy_nn_params: parameters of the PotentialEnergyNN
        th: link angles of double pendulum of shape (2, )
        th_d: link angular velocities of double pendulum of shape (2, )
        tau: link torques of double pendulum of shape (2, )
    Returns:
        th_dd: link angular accelerations of double pendulum of shape (2, )
    """
    # Compute the angular acceleration of the links
    th_dd = jnp.zeros((2,))
    # YOUR CODE HERE
    raise NotImplementedError()

    return th_dd
```

<!-- grade_id: cell-04d5fab5a1e9a1e6, points: 0, grade: false, solution: true -->
```python
def continuous_state_space_dynamics(
    mass_matrix_nn_params: FrozenDict,
    potential_energy_nn_params: FrozenDict,
    x: jnp.ndarray,
    tau: jnp.ndarray,
) -> Tuple[jnp.ndarray, jnp.ndarray]:
    """
    Compute the continuous forward dynamics of the system in state-space representation
    using the Lagrangian neural network
    Args:
        mass_matrix_nn_params: parameters of the MassMatrixNN
        potential_energy_nn_params: parameters of the PotentialEnergyNN
        x: system state of shape (4, ) consisting of the link angles and velocities
        tau: link torques of shape (2, )
    Returns:
        dx_dt: time derivative of the system state of shape (4, )
        y: system output of shape (2, ) consisting of the link angles
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return dx_dt, y
```

<!-- grade_id: cell-9290e00886679359, points: 0, grade: false, solution: true -->
```python
@jit
def discrete_forward_dynamics(
    mass_matrix_nn_params: FrozenDict,
    potential_energy_nn_params: FrozenDict,
    dt: jnp.ndarray,
    th_curr: jnp.ndarray,
    th_d_curr: jnp.ndarray,
    tau: jnp.ndarray = jnp.zeros((2,)),
) -> Tuple[jnp.ndarray, jnp.ndarray, jnp.ndarray]:
    """
    Compute the discrete forward dynamics of the system using a learned neural-network-based
    mass matrix and potential energy
    Args:
        mass_matrix_nn_params: parameters of the MassMatrixNN
        potential_energy_nn_params: parameters of the PotentialEnergyNN
        dt: time step between the current and the next state [s] of shape ( )
        th_curr: current link angles of double pendulum of shape (2, )
        th_d_curr: current link angular velocities of double pendulum of shape (2, )
        tau: link torques of double pendulum of shape (2, )
    Returns:
        th_next: link angles at the next time step of double pendulum of shape (2, )
        th_d_next: link angular velocities at the next time step of double pendulum of shape (2, )
        th_dd: link angular accelerations between current and next time step of double pendulum of shape (2, )
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return th_next, th_d_next, th_dd
```

```python

```