# RO47019: Intelligent Control Systems Practical Assignment
* Period: 2024-2025, Q4
* Course homepage: https://brightspace.tudelft.nl/d2l/home/682445
* Instructor: Cosimo Della Santina (C.DellaSantina@tudelft.nl)
* Teaching assistant: Niels Stienen (N.L.Stienen@student.tudelft.nl)
* (c) TU Delft, 2025

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE` and remove `raise NotImplementedError()` afterwards. Moreover, if you see an empty cell, please **do not** delete it, instead run that cell as you would run all other cells. Finally, please **do not** add any extra cells to this notebook or change the existing cells unless you are explicitly asked to do so.

Please fill in your name(s) and other required details below:

```python
# Please fill in your names, student numbers, netID, and emails below.
STUDENT_1_NAME = ""
STUDENT_1_STUDENT_NUMBER = ""
STUDENT_1_NETID = ""
STUDENT_1_EMAIL = ""
```

<!-- grade_id: cell-9f148ec62e0de49c, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled
assert STUDENT_1_NAME != ""
assert STUDENT_1_STUDENT_NUMBER != ""
assert STUDENT_1_NETID != ""
assert STUDENT_1_EMAIL != ""
```

### General announcements

* Do *not* share your solutions (also after the course is finished), and do *not* copy solutions from others. By submitting your solutions, you claim that you alone are responsible for this code.

* Please post your questions regarding this assignment in the correct support forum on Brightspace, this way everybody can benefit from the response. Please note that it is **not** allowed to post any code relating to solution attempts. If you do have a particular question that you want to ask directly, please use the scheduled Q&A hours to ask the TA or if not possible otherwise, send an email to the instructor or TA.

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code! These lines are just there as a reminder for you that you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore.

* This [Jupyter notebook](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. 

* Please note that the above mentioned _read-only_ protection only works in Jupyter Notebook, and it does not work if you open this notebook in another editor (e.g., VSCode, PyCharm, etc.). Therefore, we recommend that you only use Jupyter Notebook for this course. If you use any other editor, you may accidentally delete cells, modify the tests, etc., which would cause you to lose points.

* If you edit a function that is imported in another notebook, you need to **restart the kernel** of the notebook where you are using the function. Otherwise, the changes will not be effective.

* **IMPORTANT**: Please make sure that your code executes without any errors before submitting the notebook. An easy way to ensure this is to use the validation script as described in the README.

```python

```

# Task 2c.2 - Implementation of a Lagrangian Neural Network (LNN) (9p)

**Author:** Maximilian Stölzle, Jingyue Liu (J.Liu-14@tudelft.nl)


This notebook will guide you through the implementation of a Lagrangian neural network.

```python
# Reloads the python files outside of this notebook automatically
%load_ext autoreload
%autoreload 2

# import all Python modules
from distutils.util import strtobool
from functools import partial
import jax

jax.config.update("jax_platforms", "cpu")  # set default device to 'cpu'
jax.config.update("jax_enable_x64", True)  # double precision
import jax
from jax import random
from jax import numpy as jnp
import os
from pathlib import Path
import warnings

from jax_double_pendulum.integrators import rk4_step
from jax_double_pendulum.utils import normalize_link_angles

# define boolean to check if the notebook is run for the purposes of autograding
AUTOGRADING = strtobool(os.environ.get("AUTOGRADING", "false"))
# define tolerances for grading
RTOL = float(os.environ.get("RTOL", "1e-4"))  # relative tolerance
ATOL = float(os.environ.get("ATOL", "1e-7"))  # absolute tolerance
```

## Neural networks (2p)

First, please finish the implementation of the neural networks `MassMatrixNN` and `PotentialEnergyNN`. Both neural networks should consist of four layers and have the following structure:

1. A linear layer mapping from the input dimension to the hidden state dimension $n_\mathrm{hidden}$.

2. A softplus activation function.

3. A linear layer with $n_\mathrm{hidden}$ as both the input and output dimension.

4. A softplus activation function.

5. A linear layer with $n_\mathrm{hidden}$ as both the input and output dimension.

6. A softplus activation function.

7. A linear layer with $n_\mathrm{hidden}$ as both the input and output dimension.

8. A softplus activation function.

9. A linear layer mapping from $n_\mathrm{hidden}$ to the output dimension.


The hidden state dimension is set by the `num_hidden` attribute of the class to $32$. The softplus activation function is a smooth approximation of the Rectified Linear Unit (ReLU) function.

The mass matrix $M$ (sometimes also denoted as $B$) of mechanical systems has two important characteristics, which we want to enforce for our learned mass matrix as well:

- The diagonal values are always positive
- The matrix is symmetric

The symmetry is given automatically by the fact that we only learn a triangular matrix and then derive the complete mass matrix from that triangular matrix. We can induce the first characteristic by applying a softplus activation function in conjunction with a shifting of the diagonals.

<!-- grade_id: cell-eb09c2699bae145a, points: 2, grade: true, solution: false -->
```python
# DO NOT REMOVE OR MODIFY THIS CELL

# import MassMatrixNN and PotentialEnergyNN from lnn.ipynb
from ipynb.fs.full.lnn import MassMatrixNN, PotentialEnergyNN

# initialize the PRNG key at seed 0
_rng = random.PRNGKey(seed=0)

# initialize parameters of MassMatrixNN by passing a dummy input of ones through the network
_mass_matrix_nn_params = MassMatrixNN().init(_rng, jnp.ones((2,)))["params"]

# initialize parameters of PotentialEnergyNN by passing a dummy input of ones through the network
_potential_energy_nn_params = PotentialEnergyNN().init(_rng, jnp.ones((2,)))["params"]

# Some dummy system state
_th = 0 * jnp.ones((2,))

_M_nn = MassMatrixNN().apply({"params": _mass_matrix_nn_params}, _th)
print("Computed mass matrix (output from neural network):\n", _M_nn)

_M_nn_target = jnp.array([[0.54595144, -0.53372961], [-0.53372961, 0.97476694]])
print("Target mass matrix (output from neural network):\n", _M_nn_target)
if not jnp.allclose(_M_nn, _M_nn_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The output of the mass matrix neural network is not correct.")

_U = PotentialEnergyNN().apply({"params": _potential_energy_nn_params}, _th)
print("Computed potential energy:\n", _U)

# make sure that U is a scalar
if not _U.shape == ():
    warnings.warn("The PotentialEnergyNN needs to return a scalar.")

_U_target = jnp.array(-0.3061547720368191)
print("Target potential energy:\n", _U_target)
if not jnp.allclose(_U, _U_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The output of the potential energy neural network is not correct.")

```

## Lagrangian (1p)

First, please implement the two functions `kinetic_energy_fn` and `potential_energy_fn` which compute the (estimated) kinetic energy and potential energy of the system respectively:

\begin{equation}
    \mathcal{T} = \frac{1}{2} \: \dot{\theta}^\mathrm{T} \: f_\mathrm{M}(\theta, \pi_\mathrm{M}) \: \dot{\theta}, 
    \qquad 
    \mathcal{U} = f_\mathcal{U}(\theta, \pi_\mathcal{U}),
\end{equation}

where $f_\mathrm{M}(\theta, \pi_\mathrm{M})$ and $f_\mathcal{U}(\theta, \pi_\mathcal{U})$ are the two neural network forward functions respectively.

The Lagrangian $\mathcal{L}$ is then defined as

\begin{equation}
    \mathcal{L} = \mathcal{T} - \mathcal{U},
\end{equation}

and needs to be implemented into `lagrangian_fn`.

<!-- grade_id: cell-16a7f1ab8bb47b51, points: 1, grade: true, solution: false -->
```python
# DO NOT REMOVE OR MODIFY THIS CELL

# import lagrangian_fn from lnn.ipynb
from ipynb.fs.full.lnn import lagrangian_fn

# dummy link velocity
_th_d = jnp.pi * jnp.ones((2,))

_L = lagrangian_fn(
    _mass_matrix_nn_params,
    _potential_energy_nn_params,
    _th,
    _th_d,
)
print("Computed Lagrangian:", _L)

# make sure the Lagrangian is a scalar
if not _L.shape == ():
    warnings.warn("The Lagrangian needs to be a scalar.")

_L_target = jnp.array(2.542899071686838)
print("Target Lagrangian:", _L_target)
if not jnp.allclose(_L, _L_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed Lagrangian is not correct.")

```

## Continuous forward dynamics (4p)

Next, please take the various partial derivatives of the Lagrangian to compute the dynamical matrices (i.e. the components of the equations of motion) in the function `dynamical_matrices`

\begin{equation}
    M = \frac{\partial^2 \mathcal{L}}{\partial \dot{\theta}^2},
    \qquad
    C_{ij} = \sum_{k=1}^{2} \frac{1}{2} \left ( \frac{\partial M_{ij}}{\partial \theta_k}  + \frac{\partial M_{ik}}{\partial \theta_j} -\frac{\partial M_{jk}}{\partial \theta_i} \right ) \, \dot{\theta}_k \: \forall \: i, j \in \{1, 2\},
    \qquad
    G = \frac{\partial \mathcal{U}}{\partial \theta}.
\end{equation}

Make sure that before you evaluate the Lagrangian, you normalize the link angles $\theta$ to the interval $[-\pi, \pi]$. If we were not to implement this normalization, we would (likely) get a different behaviour for $\theta = a$ or $\theta = a + 2 \: \pi$, which is not desired. This would also drastically increase the amount of training data we would require.

Then, please compute the continuous forward dynamics (i.e. determining $\ddot{\theta}$ for a given state $(\theta, \dot{\theta})$ and input $\tau$) in the function `continuous_forward_dynamics`. For this implementation, you can take inspiration from the analytical solution in `jax_double_pendulum.dynamics.continuous_forward_dynamics`.

<!-- grade_id: cell-3604fa447bb6bff9, points: 4, grade: true, solution: false -->
```python
# DO NOT REMOVE OR MODIFY THIS CELL

# import dynamical_matrices from lnn.ipynb
from ipynb.fs.full.lnn import dynamical_matrices, continuous_forward_dynamics

_M, _C, _G = dynamical_matrices(
    _mass_matrix_nn_params,
    _potential_energy_nn_params,
    _th,
    _th_d,
)

_M_target = jnp.array([[0.54595144, -0.53372961], [-0.53372961, 0.97476694]])
_C_target = jnp.array([[-0.00735404, 0.17485702], [0.09947973, -0.12300335]])
_G_target = jnp.array([0.03959358, -0.05480955])

print("Computed M:\n", _M)
if not _M.shape == (2, 2):
    warnings.warn("The mass matrix needs to have the shape (2, 2)")

# the mass matrix estimated by the neural network should be the same as d2L_d2thd
if not jnp.allclose(_M_nn, _M, rtol=RTOL, atol=ATOL):
    warnings.warn(
        "The mass matrix computed from the Lagrangian needs to match the output of the neural network."
    )

print("Target M:\n", _M_target)
if not jnp.allclose(_M, _M_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed mass matrix is not correct.")

print("Computed C:\n", _C)
if not _C.shape == (2, 2):
    warnings.warn("The Coriolis matrix needs to have the shape (2, 2)")

print("Target C:\n", _C_target)
if not jnp.allclose(_C, _C_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed coriolis matrix is not correct.")

print("Computed G:\n", _G)
if not _G.shape == (2, ):
    warnings.warn("The gravity vector needs to have the shape (2, )")

print("Target G:\n", _G_target)
if not jnp.allclose(_G, _G_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed gravity vector is not correct.")

# zero applied torque
_tau = jnp.ones((2,))

_th_dd = continuous_forward_dynamics(
    _mass_matrix_nn_params,
    _potential_energy_nn_params,
    _th,
    _th_d,
    _tau,
)

_th_dd_target = jnp.array([4.14726048, 3.42874464])

print("Computed th_dd:\n", _th_dd)
print("Target th_dd:\n", _th_dd_target)
if not jnp.allclose(_th_dd, _th_dd_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed angular link acceleration is not correct.")

```

## State-space dynamics (1p)

Please implement dynamics in a nonlinear state-space representation into the function `continuous_state_space_dynamics`. The state-space representation needs to have the following structure

\begin{equation}
    \frac{\mathrm{d}x}{\mathrm{d}t} = f(x, \tau),\\
    y = g(x, \tau),
\end{equation}

where $x \in \mathbb{R}^4$ is the state of the system. For implementing this function, you can take inspiration from `jax_double_pendulum.dynamics.continuous_state_space_dynamics`.

Having the EoM in this structure will allow us in the next step to use standard methods such as [RK4](https://en.wikipedia.org/wiki/Runge–Kutta_methods) to integrate the ODE in time.

<!-- grade_id: cell-82477ae2e15f5b3b, points: 1, grade: true, solution: false -->
```python
# DO NOT REMOVE OR MODIFY THIS CELL

# import continuous_state_space_dynamics from lnn.ipynb
from ipynb.fs.full.lnn import continuous_state_space_dynamics

# define a dummy state
_x = jnp.concatenate([_th, _th_d])

_dx_dt, _y = continuous_state_space_dynamics(
    _mass_matrix_nn_params,
    _potential_energy_nn_params,
    _x,
    _tau,
)

_dx_dt_target = jnp.array([3.14159265, 3.14159265, 4.14726048, 3.42874464])
_y_target = jnp.array([[0.0, 0.0]])

print("Computed dx_dt:\n", _dx_dt)
print("Target dx_dt:\n", _dx_dt_target)
if not jnp.allclose(_dx_dt, _dx_dt_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed time derivative of the state is not correct.")

print("Computed y:", _y, "target y:", _y_target)
if not jnp.allclose(_y, _y_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed system output is not correct.")

```

## Discrete-time forward dynamics (1p)

Now, we move on to implementing the discrete-time forward dynamics into `discrete_forward_dynamics`.
Please use the `jax_double_pendulum.integrators.rk4_step` function to integrate the continuous state-space dynamics for one time-step of duration `dt`. The `rk4_step` expects as its first argument an `ode_fun`, which needs to conform to the following syntax: `rk4_step(x) -> dx_dt` (i.e. a mapping from the state $x$ to its time-derivative $\dot{x}$). For implementing this function, you can take inspiration from `jax_double_pendulum.dynamics.discrete_forward_dynamics`.

<!-- grade_id: cell-d635ffa07cc2e768, points: 1, grade: true, solution: false -->
```python
# DO NOT REMOVE OR MODIFY THIS CELL

# import discrete_forward_dynamics from lnn.ipynb
from ipynb.fs.full.lnn import discrete_forward_dynamics

_dt = 1e-2  # time-step [s]
_th_next, _th_d_next, _th_dd = discrete_forward_dynamics(
    _mass_matrix_nn_params,
    _potential_energy_nn_params,
    _dt,
    _th,
    _th_d,
    _tau,
)

_th_next_target = jnp.array([0.03162163, 0.0315865])
_th_d_next_target = jnp.array([3.18256705, 3.17561977])
_th_dd_target = jnp.array([4.14726048, 3.42874464])

print("Computed th_next:\n", _th_next)
print("Target th_next:\n", _th_next_target)
if not jnp.allclose(_th_next, _th_next_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed next link angles are not correct.")

print("Computed th_d_next:\n", _th_d_next)
print("Target th_d_next:\n", _th_d_next_target)
if not jnp.allclose(_th_d_next, _th_d_next_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed next angular link velocities are not correct.")

print("Computed th_dd:\n", _th_dd)
print("Target th_dd:\n", _th_dd_target)
if not jnp.allclose(_th_dd, _th_dd_target, rtol=RTOL, atol=ATOL):
    warnings.warn("The computed angular link accelerations are not correct.")


```